# CS556 Data Security and Privacy
## Homework 3 Differential Privacy I
### Due date & time: 11:59pm on March 4, 2025.

Filling the code blocks and answer the questions.

```python
# Load the data and libraries
import pandas as pd
import numpy as np
from scipy import stats
import matplotlib.pyplot as plt
plt.style.use('seaborn-whitegrid')
# plt.style.use('seaborn-v0_8-whitegrid')
from unittest.mock import patch


adult = pd.read_csv('https://raw.githubusercontent.com/zealscott/CS556/main/adult_data.csv')
adult = adult.dropna()
```

```python
# varify the downloaded adult dataset
adult.head()
```

Output:
```
Name        DOB          SSN    Zip  Age         Workclass  \
0  Karrie Trusslove   9/7/1967  732-14-6110  64152   39         State-gov   
1  Brandise Tripony   6/7/1988  150-19-2766  61523   50  Self-emp-not-inc   
2     Brenn McNeely   8/6/1991  725-59-9860  95668   38           Private   
3       Dorry Poter   4/6/2009  659-57-4974  25503   53           Private   
4       Dick Honnan  9/16/1951  220-93-3811  75387   28           Private   

   fnlwgt  Education  Education-Num      Marital Status         Occupation  \
0   77516  Bachelors             13       Never-married       Adm-clerical   
1   83311  Bachelors             13  Married-civ-spouse    Exec-managerial   
2  215646    HS-grad              9            Divorced  Handlers-cleaners   
3  234721       11th              7  Married-civ-spouse  Handlers-cleaners   
4  338409  Bachelors             13  Married-civ-spouse     Prof-specialty   

    Relationship   Race     Sex  Capital Gain  Capital Loss  Hours per week  \
0  N
```

## Question 1: Warm-up with Pandas (6 points)

Pandas is a very powerful tabular data processing tool.

[Query 1] (2 points) Write code to calculate how many rows in the dataset have the value "Bachelors" in the "Education" attribute, and "Age" at most 30?

[Query 2] (2 points) Write code to generate the histogram over the "Education" attribute. The return should be a dictionary.

[Query 3] (2 points) Write code to calculate the range of "Age". The return should be two integers.

Hint: Some Pandas functions could be very useful, e.g. `value_counts()`.

<!-- grade_id: cell-975eb03979d78eaf, points: 0, grade: false, solution: true -->
```python
def Query1():
    # YOUR CODE HERE
    return None

def Query2():
    # YOUR CODE HERE
    return None

def Query3():
    # YOUR CODE HERE
    min, max = -1, -1
    return min, max

print(f'The number of rows with value "Bachelors" in the "Education" attribute, and "Age" at most 30 is {Query1()}.')
print(f'The histogram counts of "Education" attribute is: {Query2()}.')
max_age, min_age = Query3()
print(f'The max age is {max_age}; the min age is {min_age}')

```

Output:
```
The number of rows with value "Bachelors" in the "Education" attribute, and "Age" at most 30 is None.
The histogram counts of "Education" attribute is: None.
The max age is -1; the min age is -1
```

## Question 2: Sensitivity of queries (10 points)

Answer the following questions:
- (3 points) What is the sensitivity of `Query1`, and why?
- (3 points) Assume the domain of "Education" is known beforehand. What is the sensitivity of `Query2`, and why?
- (4 points) What is the sensitivity of "the average age"? Do we need additional conditions/assumptions to calculate it?

<!-- grade_id: cell-fb1c7b0533f933f7, points: 5, grade: true, solution: true -->
YOUR ANSWER HERE

## Question 3: Laplace mechanism (14 points)

- (4 points) Implement the `laplace_mech` to produce a differentially private answer given an input value `v`, `sensitivity` and privacy budget `epsilon`
- (4 points) Use your implemented `laplace_mech` to generate private answers for Query1 with given $\epsilon$.
- (6 points) Use your implemented `laplace_mech` to generate private answers for Query2 with given $\epsilon$.

<!-- grade_id: cell-80d3c108ba0f75d5, points: 0, grade: false, solution: true -->
```python
def laplace_mech(v, sensitivity, epsilon):
    # YOUR CODE HERE
    return None

def dp_query1(epsilon):
    # YOUR CODE HERE
    return None

def dp_query2(epsilon):
    # YOUR CODE HERE
    return None

dp_query1(0.1)
```

```python
dp_query2(0.1)
```

<!-- grade_id: cell-93eab43d27806309, points: 10, grade: true, solution: false -->
```python
# TEST CASE for question 3
dp_results = [dp_query1(0.1) for _ in range(100)]
spec = [np.random.laplace(loc=1475, scale=1/0.1) for _ in range(100)]
assert stats.wasserstein_distance(dp_results, spec) < 10
```

```python
# compare the results of query 2 and its dp version

ground_truth = Query2()
dp_result = dp_query2(0.1)

fig, (ax1, ax2) = plt.subplots(1, 2)
ax1.bar(ground_truth.keys(), ground_truth.values(), alpha=0.3)
ax2.bar(dp_result.keys(), dp_result.values(), alpha=0.3)
ax1.tick_params(axis='x', labelrotation=90)
ax2.tick_params(axis='x', labelrotation=90)
plt.show()
```

## Question 4: Analyse the utility of Laplace mechanism (20 points)

Showing the relationship between the query result errors and different values privacy budget epsilon is always required to show the effectiveness or utility of a DP mechanism.

Consider the (private) query 2 result as a vector. The error can be measured by L1 or L2 distance between the ground truth results and privatized results.

- (2 points) implement the L1 distance function, with `ground_truth` and `privatized_result` as input and output L1 distance.
- (2 points) implement the L2 distance function, with `ground_truth` and `privatized_result` as input and output L1 distance


```python
def l1_distance(ground_truth, privatized_result):
    # YOUR CODE HERE
    return None


def l2_distance(ground_truth, privatized_result):
    # YOUR CODE HERE
    return None
```

```python
q2_result = Query2()
keys = q2_result.keys()
privatized_result = dp_query2(0.1)
ground_truth = np.array([q2_result[k] for k in keys])
privatized_result = np.array([privatized_result[k] for k in keys])
print(f"L1 error: {l1_distance(ground_truth, privatized_result)}")
print(f"L2 error: {l2_distance(ground_truth, privatized_result)}")
```

- (6 points) Given a list privacy budget epsilons, implement a function that repeats the `dp_query2` 100 times for each `epsilon`,
compute the average L1 error of the privatized result, and plot it in a graph.



<!-- grade_id: cell-da0c198a1cf9a866, points: 0, grade: false, solution: true -->
```python
epsilons = [0.1, 0.2, 0.4, 0.6, 0.8, 1, 1.2, 1.4, 1.6, 1.8, 2.0]
keys = ['HS-grad', 'Some-college', 'Bachelors', 'Masters', 'Assoc-voc', '11th', 'Assoc-acdm', '10th', '7th-8th', 'Prof-school', '9th', '12th', 'Doctorate', '5th-6th', '1st-4th', 'Preschool']

def analyse_dp_q2_l1_error():
    # YOUR CODE HERE
    # plot a figure with epsilons as x axis and L1 error as y axis
    raise NotImplementedError

analyse_dp_q2_l1_error()
```

- (2 points) Given a list privacy budget epsilons, implement a function that repeats the `dp_query2` 100 times for each `epsilon`,
compute the average L2 error of the privatized result, and plot it in a graph.



<!-- grade_id: cell-eda6bc27840a9067, points: 10, grade: true, solution: false -->
```python
def analyse_dp_q2_l2_error():
    # YOUR CODE HERE
    # plot a figure with epsilons as x axis and L1 error as y axis
    raise NotImplementedError

analyse_dp_q2_l2_error()
```

- (8 points) Compare the result of L1 errors and L2 errors. Describe the relations between the errors and different `epsilon` values.
Which type of error (L1 or L2) can better align with the Laplace mechanism?

YOUR ANSWER HERE

## Question 6: Clipping (20 points)
It is very common that some query's sensitivity is not bounded or too large.
To handle such case, clipping is usually a solution.
By clipping, we mean that given a pair of thresholds $(\theta_{\min}, \theta_{\max})$,
we pre-processing the dataset so that $\hat{v} = \max \{ \min \{v, \theta_{\max} \}, \theta_{\min}\}$.

- (2 points) Write Query 4 to calculate the average of `Capital Gain`


```python
def Query4():
    # YOUR CODE HERE
    return None

Query4()
```

<!-- grade_id: cell-075013c43cc4cc1d, points: 10, grade: true, solution: true -->
- (2 points) We can assume that capital gain should be non-negative numbers. Implement the function that returns
differential private answer to Query 4 given $\theta_{\max}$ and `epsilon`.


```python
def dp_query4(theta, epsilon):
    # YOUR CODE HERE
    return None

dp_query4(10000, 1)
```

- (8 points) For private Qeury 4, empirically find the best choice $\theta_{\max}$ from the following set of values `[100, 500, 1000, 5000, 10000, 50000, 100000, 200000, 500000]` given $\epsilon = 1$.

```python
thetas = [100, 500, 1000, 5000, 10000, 50000, 100000, 200000, 500000]
eps = 1

def find_opt_theta():
    # YOUR CODE HERE
    opt = -1
    return opt

print(f"The optimal threshold is {find_opt_theta()}")
```

- (6 points) Summary the finding of different threshold values, and explain why.

YOUR ANSWER HERE

## Question 5: Noisy-min (30 points)
- (5 points) Implement a function to return the number of records having `Age` greater than 50 within different education groups.
(Hint: pandas `groupby` function can make it easier)


```python
def Query5():
    # YOUR CODE HERE
    return None

Query5()
```

 - (10 points) Implement a differential private function to return the `Education` group with the smallest number of records with `Age` > 50.
Clip the `Age` to be in [0, 100]. Also, explain how to set the sensitivity.


```python
def dp_query5(epsilon):
    # YOUR CODE HERE
    return None

print(f"The groun with smallest number of older than 50 is {dp_query5(epsilon=1)}")
```

- (15 points) Is it still satisfying DP if the function also returns the randomized count of the Education group? Why?